Before you begin, execute this cell to import numpy and packages from the D-Wave Ocean suite, and all necessary functions for the gate-model framework you are going to use, whether that is the Forest SDK or Qiskit. In the case of Forest SDK, it also starts the qvm and quilc servers.

```python
%run -i "assignment_helper.py"
%matplotlib inline
```

Output:
```
RuntimeError                              Traceback (most recent call last)
File ~/QML/qml-mooc/coding_assignments/assignment_helper.py:156, in <module>
    154     is_dwave = False
    155 if not (is_qiskit or is_forest):
--> 156     raise RuntimeError("No quantum computing framework available!")
    157 if not is_dwave:
    158     raise RuntimeError("D-Wave Ocean is not available!")

RuntimeError: No quantum computing framework available!
```

```python
import sys
!{sys.executable} -m pip install matplotlib
```

Output:
```
Defaulting to user installation because normal site-packages is not writeable
Collecting matplotlib
  Downloading matplotlib-3.5.2-cp310-cp310-manylinux_2_17_x86_64.manylinux2014_x86_64.whl (11.9 MB)
[2K     [38;2;114;156;31m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m11.9/11.9 MB[0m [31m10.3 MB/s[0m eta [36m0:00:00[0mm eta [36m0:00:01[0m[36m0:00:01[0m
[?25hRequirement already satisfied: pyparsing>=2.2.1 in /usr/lib/python3/dist-packages (from matplotlib) (2.4.7)
Requirement already satisfied: python-dateutil>=2.7 in /home/shubham/.local/lib/python3.10/site-packages (from matplotlib) (2.8.2)
Requirement already satisfied: pillow>=6.2.0 in /usr/lib/python3/dist-packages (from matplotlib) (9.0.1)
Requirement already satisfied: packaging>=20.0 in /home/shubham/.local/lib/python3.10/site-packages (from matplotlib) (21.3)
Collecting kiwisolver>=1.0.1
  Downloading kiwisolver-1.4.3-cp310-cp310-manylinux_2_12_x86_64.manylinux2010_x86_64.whl (1.6 MB)
[2K     [38;2;114;156;31
```

```python
import numpy as np
```

# Classical probability distributions

**Exercise 1** (1 point). Recall that in classical con flipping, get heads with probability $P(X=0) = p_0$ and tails with $P(X=1) = p_1$ for each toss of the coin, where $p_i\geq 0$ for all $i$, and the probabilities sum to one: $\sum_i p_i = 1$. Create a sample with a 1000 data points using numpy, with a probability of getting tails being 0.3. This is the parameter that the `binomial` function takes. Store the outcome in an array called `x_data`.

```python
n_samples = 1000
###
### YOUR CODE HERE
x_data = np.empty((1000,))
p_1 = 0.3
def binomial(pr):
    for i,num in enumerate(range(n_samples)):
    # 0-> heads, 1-> tails
        x_data[i] = np.random.choice([0,1], p=[1-pr,pr])

   
    
binomial(p_1)

###

```

<!-- grade_id: exercise1, points: 1, grade: true, solution: false -->
```python
assert isinstance(x_data, np.ndarray)
assert abs(p_1-x_data.sum()/n_samples) < 0.05
```

**Exercise 2** (1 point). As you recall, we may also write the probability distribution as a stochastic vector $\vec{p} = \begin{bmatrix} p_0 \\ p_1 \end{bmatrix}$. The normalization constraint on the probability distribution says that the norm of the vector is restricted to one in the $l_1$ norm. In other words, $||\vec{p}||_1 = \sum_i |p_i| = 1$. This would be the unit circle in the $l_1$ norm, but since $p_i\geq 0$, we are restricted to a quarter of the unit circle, just as we plotted above. Write a function that checks whether a given two-dimensional vector is a stochastic vector. That is, it should return `True` if all elements are positive and the 1-norm is approximately one, and it should return `False` otherwise. The input of the function is a numpy array.

```python
def is_stochastic_vector(p: np.array):
    ###
    ### YOUR CODE HERE
    positive = not np.any(p<0)
    summ = np.sum(p)
    sum_bool = False
    if 1-summ < 0.1 and 1-summ >= 0:
        sum_bool = True    
    return positive and sum_bool   




```

<!-- grade_id: exercise2, points: 1, grade: true, solution: false -->
```python
assert not is_stochastic_vector(np.array([0.2, 0.3]))
assert not is_stochastic_vector(np.array([-0.2, 0.7]))
assert is_stochastic_vector(np.array([0.2, 0.8]))
```

**Exercise 3** (1 point). The probability of heads is just the first element in the $\vec{p}$ and we can use a projection to extract it. For the first element of the stochastic vector, the projection is described by the matrix $\begin{bmatrix} 1 & 0\\0 & 0\end{bmatrix}$. Write a function that performs this projection on a two-element vector described by a numpy array. Your output after the projection is also a two-element vector.

```python
def project_to_first_basis_vector(p: np.array):
    ###
    ### YOUR CODE HERE
    heads_prob = np.dot(p,np.array([[1,0],[0,0]]))
    return heads_prob
    ###
# assert np.alltrue(project_to_first_basis_vector(np.array([0.2, 0.3])) == np.array([0.2, 0.]))
# assert np.alltrue(project_to_first_basis_vector(np.array([1., 0.])) == np.array([1., 0.]))
```

<!-- grade_id: exercise3, points: 1, grade: true, solution: false -->
```python
assert np.alltrue(project_to_first_basis_vector(np.array([0.2, 0.3])) == np.array([0.2, 0.]))
assert np.alltrue(project_to_first_basis_vector(np.array([1., 0.])) == np.array([1., 0.]))
```

**Exercise 4** (1 point). The projection operators introduce some linear algebra to working with probability distributions. We can also use linear algebra to transform one probability distribution to another. A left *stochastic matrix* will map stochastic vectors to stochastic vectors when multiplied from the left: its columns add up to one. Write a function that takes a matrix and a vector as input arguments (both are numpy arrays), checks whether the vector is a stochastic vector and whether the matrix is left stochastic. If they are, return the matrix applied to the vector, otherwise raise a `ValueError`. You can call the function `is_stochastic_vector` that you defined above.

```python
def apply_stochastic_matrix(p: np.array, M: np.array):
    """Apply the matrix M to the vector p, but only if
    p is a stochastic vector and M is a left stochastic
    matrix. Otherwise raise a ValueError.
    """
    ###
    ### YOUR CODE HERE
    p_check = is_stochastic_vector(p)
    if(p_check):
        transformed = np.dot(M,p)
        M_check = is_stochastic_vector(transformed)
        if M_check:
            return transformed
        else:
            raise ValueError
    else:
        raise ValueError    
    ###

```

<!-- grade_id: exercise4, points: 1, grade: true, solution: false -->
```python
p = np.array([[.5], [.5]])
M = np.array([[0.7, 0.6], [0.3, 0.4]])
assert abs(np.linalg.norm(apply_stochastic_matrix(p, M), ord=1)-1) < 0.01
M = np.array([[0.7, 0.6], [0.3, 0.5]])
try:
    apply_stochastic_matrix(p, M)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

**Exercise 5** (1 point). Create a left stochastic matrix in a variable called `M` that transforms the uniform distribution $\vec{p}= \begin{bmatrix} 0.5 \\ 0.5 \end{bmatrix}$ to $\begin{bmatrix} 0.6 \\ 0.4 \end{bmatrix}$. `M` should be a two-dimensional numpy array.

```python
###
### YOUR CODE HERE
a = np.random.rand()
b = 1.2-a
c = random.rand()
d = 0.8-c
M = np.array([a,b,c,d]).reshape(2,2)
###

```

<!-- grade_id: exercise5, points: 1, grade: true, solution: false -->
```python
assert np.allclose(M.dot(np.array([0.5, 0.5])), np.array([0.6, 0.4]))
```

**Exercise 6** (1 point). Calculate the entropy of this distribution $\begin{bmatrix} 0.6 \\ 0.4 \end{bmatrix}$ in a variable called `S`.

```python
###
### YOUR CODE HERE
dist = np.array([0.6,0.4]).reshape(2,1)
dist_log = np.log2(dist)

S = np.dot(dist_log.reshape(1,2), dist)
S = S*-1
float(S)
###

```

Output:
```
0.9709505944546686
```

<!-- grade_id: exercise6, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

# Quantum states

**Exercise 7** (1 point). A quantum state is a probability distribution. A qubit state is a distribution over two values, similar to the coin flipping in the classical state. A major difference is that the entries are complex numbers and the normalization is in the $l_2$ norm. Create a function similar to `is_stochastic_vector` that checks whether a vector is a valid quantum state. The input is a numpy array and the output should be boolean.

```python
def is_quantum_state(psi: np.array):
    ###
    ### YOUR CODE HERE
    return np.abs(np.sum(np.square(np.abs(psi)))-1)<=0.1


```

<!-- grade_id: exercise7, points: 1, grade: true, solution: false -->
```python
assert is_quantum_state(np.array([1/np.sqrt(2), 1/np.sqrt(2)]))
assert is_quantum_state(np.array([-1/np.sqrt(2), 1/np.sqrt(2)]))
assert is_quantum_state(np.array([-1/3, 2*np.sqrt(2)/3]))
assert is_quantum_state(np.array([-1j/3, 2*np.sqrt(2)/3]))
assert not is_quantum_state(np.array([0.2, 0.8]))
```

**Exercise 8** (1 point). While working with numpy arrays is convenient, it is better to use a framework designed for quantum computing, since it often allows us to execute a circuit directly on a quantum computer. In your preferred framework, implement a circuit of a single qubit with no operation on it. You should create it in an object called `circuit`. Do not add a measurement. The evaluation will automatically branch according to which framework you chose.

```python
###
### YOUR CODE HERE
from qiskit import QuantumCircuit as qci
circuit = qci(1,1)

###

```

<!-- grade_id: exercise8, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert abs(amplitudes[0]-1.0) < 0.01
```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [49], in <cell line: 1>()
----> 1 amplitudes = get_amplitudes(circuit)
      2 assert abs(amplitudes[0]-1.0) < 0.01

NameError: name 'get_amplitudes' is not defined
```

**Exercise 9** (1 point). In the execution branching above, you see that we use the wavefunction simulator. This allows us to use the probability amplitudes as usual numpy arrays, as you can see above. If we ran the circuit on an actual quantum device, we would not be able to inspect the wavefunction, but we would have to rely on the statistics of measurements to understand what is happening in the circuit.

Create a circuit in your preferred framework that creates an equal superposition in a qubit using a Hadamard gate. Again, the name of the object should be `circuit`. The evaluation will be based on measurement statistics. In this case, you should explicitly specify the measurement on the qubit

```python

from qiskit import QuantumCircuit
circuit = QuantumCircuit(1, 1)
circuit.h([0])  
circuit.measure([0], [0])

```

Output:
```
<qiskit.circuit.instructionset.InstructionSet at 0x7faac2eb3740>
```

```python
from qiskit.providers.aer import AerSimulator
Aer = AerSimulator()  # make new simulator object
job = Aer.run(circuit, shots=100)      # run the experiment
result = job.result()  # get the results
   # interpret the results as a "counts" dictionary

```

```python
def get_counts(circuit):
    print(result.get_counts())
    return result.get_counts() 
```

<!-- grade_id: exercise9, points: 1, grade: true, solution: false -->
```python
counts = get_counts(circuit)
assert abs(counts['0']/100-.5) < 0.2
```

Output:
```
{'0': 51, '1': 49}
```

**Exercise 10** (1 point). If you plotted the state before measurement on the Bloch sphere, it would have been on the equator halfway between the $|0\rangle$ and $|1\rangle$ states, and the tip of the X axis. If you apply the Hadamard on the $|1\rangle$, it would have been the point on the opposite and of the X axis, since the resulting superposition would have had a -1 amplitude for $|1\rangle$. The measurement statistics, however, would be identical. The negative sign plays a role in interference: for instance, applying a Hadamard again, would take you back to $|1\rangle$. Create the superposition after applying the Hadamard gate on $|1\rangle$. We will verify whether it picked up the phase. Do not include a measurement, since we will inspect the wavefunction.

```python
###
### YOUR CODE HERE

from qiskit import Aer, execute
circuit = QuantumCircuit(1,1)
circuit.x([0])
circuit.h([0])

###
def get_amplitudes(circuit):
    if isinstance(circuit, qiskit.circuit.quantumcircuit.QuantumCircuit):
        backend = Aer.get_backend('statevector_simulator')
        job = execute(circuit, backend)
        amplitudes = job.result().get_statevector(circuit)
    return amplitudes
```

<!-- grade_id: exercise10, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert abs(amplitudes[1]+np.sqrt(2)/2) < 0.01
```

# More qubits and entanglement

**Exercise 11** (1 point). To get a sense of multiqubit states, it is important to be confident with the tensor product operation. Create a function that returns the four basis vectors, $|00\rangle$, $|01\rangle$, $|10\rangle$, and $|11\rangle$, of the tensor product space $\mathbb{C}^2\otimes\mathbb{C}^2$. The order in which they appear does not matter. The return value should be a list of four numpy arrays.

```python
def create_canonical_basis():
    
    ###
    ### YOUR CODE HERE
    basis_R4 = []
    basis_R2 = [np.array([1,0]), np.array([0,1])]
    for cat1 in basis_R2:
        for cat2 in basis_R2:
            basis_R4.append(np.kron(cat1,cat2))
    return basis_R4
    ###

```

<!-- grade_id: exercise11, points: 1, grade: true, solution: false -->
```python
basis = create_canonical_basis()
assert len(basis) == 4
if basis[0].shape != (4, ):
    basis = [basis_vector.reshape((4, )) for basis_vector in basis]
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**Exercise 12** (1 point). A generic product state has the form  $\begin{bmatrix}a_0b_0\\ a_0b_1\\ a_1b_0\\ a_1b_1\end{bmatrix}=a_0b_0|00\rangle + a_0b_1|01\rangle + a_1b_0|10\rangle + a_1b_1|11\rangle$ on $\mathbb{C}^2\otimes\mathbb{C}^2$, but not all. We can use the basis vectors to form vectors in the space that do not have a product structure. These are entangled states that show strong correlations. Entanglement is an important resource in quantum computing and being able to create a circuit that generates an entangled state is critical. Implement a circuit in your preferred framework to create the $|\phi^-\rangle = \frac{1}{\sqrt{2}}(|00\rangle-|11\rangle)$ state, that is, almost the same as the $|\phi^+\rangle$ state, but with the opposite sign of the probability amplitude of $|11\rangle$. Do not include a measurement, as we will verify the state with the wavefunction simulator

```python
###
### YOUR CODE HERE
circuit = QuantumCircuit(2,2)
circuit.x(0)
circuit.h(0)
circuit.cx(0,1)
###

```

Output:
```
<qiskit.circuit.instructionset.InstructionSet at 0x7faa77a2c680>
```

<!-- grade_id: exercise12, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(np.array([np.sqrt(2)/2, 0, 0, -np.sqrt(2)/2]), amplitudes)
```

```python

```